# Exercise 4
Sept 16, 2019. Due Sept 18, 9am.

Enter the names, and describe the contributions of anybody besides yourself who contributed/collaborated. 

Vihara Jayaweera

## Functions

Write functions, called `my_mean()`, `my_std()`, and `my_var()` that calculate the mean, standard deviation, and variance according to the equations in exercise 3. 

You are welcome to re-use your code from exercise 3, but do not use the `numpy` `std()`, `var()` and `mean()` functions to calculate these quantities (`numpy.sum()` is ok!).

The functions should work properly with an array as input, but must not produce errors if called with an array of size 0, or if the input argument is something other than an array. In either of these cases, they should return 0 after printing the message: invalid argument


<!-- grade_id: cell-c3fe069be974d3ab, points: 0, grade: false, solution: true -->
```python
#YOUR CODE HERE
import numpy as np 

def my_mean(arr):
    if (isinstance(arr,np.ndarray) == False or arr.size == 0):
        print('invalid argument')
        return 0
    return (1/arr.size)*arr.sum()
    
def my_std(arr):
    if(isinstance(arr,np.ndarray) == False or arr.size == 0):
        print('invalid argument')
        return 0
    return np.sqrt(my_var(arr))

def my_var(arr):
    if (isinstance(arr,np.ndarray) == False or arr.size == 0 or arr.size == 1):
        print('invalid argument')
        return 0
    return (1/(arr.size-1))*((arr-my_mean(arr))**2).sum()
```

<!-- grade_id: cell-959434390e5d0c2c, points: 1, grade: true, solution: false -->
```python
# test for functions:
```

<!-- grade_id: cell-54f5ea8c5c135b7c, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-b7a4e21e897cbec2, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-a597d46d056f4c59, points: 1, grade: true, solution: false -->
```python
# test return value of my_mean()
```

<!-- grade_id: cell-d7f1138ac679272e, points: 1, grade: true, solution: false -->
```python
# test return value of my_var()
```

<!-- grade_id: cell-b8676612818d26ac, points: 1, grade: true, solution: false -->
```python
# test return value of my_std()
### BEGIN HIDDEN TEST
a = np.array([1.,3,6,9,2,3,4,73.4,27.6,95.3,225.4])
assert_almost_equal(a.std(ddof=1), my_std(a),6)
a = np.array([2,3,4,5,6])
assert_almost_equal(a.std(ddof=1), my_std(a),6)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-5-60f8e4917c02> in <module>()
      2 ### BEGIN HIDDEN TEST
      3 a = np.array([1.,3,6,9,2,3,4,73.4,27.6,95.3,225.4])
----> 4 assert_almost_equal(a.std(ddof=1), my_std(a),6)
      5 a = np.array([2,3,4,5,6])
      6 assert_almost_equal(a.std(ddof=1), my_std(a),6)

NameError: name 'assert_almost_equal' is not defined
```

<!-- grade_id: cell-c066712fc7e87050, points: 1, grade: true, solution: false -->
```python
# test invalid inputs of my_mean(), my_var() and my_std() (not all tests are shown here)
# test an empty array:
b = np.array([])
assert(my_mean(b) == 0)
assert(my_var(b) == 0)
assert(my_std(b) == 0)
# test a string
assert(my_mean("a string") == 0)
assert(my_var("a string") == 0)
assert(my_std("a string") == 0)

```

## Modules
Now take your three functions and, using a text editor like idle or nano, put them into a file in this exercise directory called `my_stats.py`. On the lab computers, the full path of this file should be /home2/__username__/p210-nbgrader/Exercise_4/my_stats.py

In the cell below you should test that you can import your module and use the functions within. If you find unexpected errors, try selecting Kernel->Restart and Run All


<!-- grade_id: cell-49ef65a0524527ce, points: 0, grade: false, solution: true -->
```python
#YOUR CODE HERE
import my_stats

```

<!-- grade_id: cell-d7682d505cf88984, points: 1, grade: true, solution: false -->
```python
# test importing my_stats.py
```

<!-- grade_id: cell-26322dfb8f518a0f, points: 1, grade: true, solution: false -->
```python
# test using functions in my_stats.py
```

## Recursion
The Fibonacci sequence: 1, 1, 2, 3, 5, 8, 13, ... is defined as:
$$F_n = F_{n-1} + F_{n-2} $$
with $F_2 = 1, F_1 =1$.

Write a function, `fibonacci(n)` to calculate the n'th number of the sequence $F_n$. Use the example from the lecture (`fact(n)`) as an inspiration. Your function should check to ensure that the argument is an integer, and return 0 if it is not. It should also return 0 in case an integer less than 1 is passed in. Depending on how the calculation is done, it may take a long time for large values of `n`, don't try to calculate values for n greater than about 30.

<!-- grade_id: cell-82b13f02d95d8e60, points: 0, grade: false, solution: true -->
```python
#YOUR CODE HERE
def fibonacci(n):
    if(not isinstance(n, int) or n < 1 or n > 30):
        return 0
    elif( n == 1 or n == 2):
        return 1
    return (fibonacci(n-1) + fibonacci(n-2))

```

<!-- grade_id: cell-92acf93c3f204c91, points: 1, grade: true, solution: false -->
```python
# test return vals of fibonacci(n)
```

<!-- grade_id: cell-e4fb462cc0861585, points: 1, grade: true, solution: false -->
```python
# test for invalid input
```

## Optional arguments

In the cell below, write a new variance function that behaves similarly to numpy's variance function, meaning that it can be called either with just one argument, in which case it will calculate the variance assuming ddof = 0, or it can be passed a second parameter that will behave as `numpy.var()`'s ddof. Call this new function `my_var2()`. 

<!-- grade_id: cell-47b85f63296ab753, points: 0, grade: false, solution: true -->
```python
#YOUR CODE HERE
def my_var2(arr,ddof = 0):
    return 1/((arr.size-ddof))*((arr-my_mean(arr))**2).sum()
       
```

<!-- grade_id: cell-ceadbd4e5bd15cb1, points: 1, grade: true, solution: false -->
```python
# test my_var2 can be called
my_var2(np.array([1,2,3]),1)
```

<!-- grade_id: cell-e5f8792c6eb0de28, points: 1, grade: true, solution: false -->
```python
# test return value of my_var2()
```

## Weighted Mean

Write a function called `wmean()` that will calculate the weighted mean of a distribution. The function should take two one-dimensional arrays as arguments. The first array can be thought of as the coordinates on the x-axis of a histogram, and the second array can be thought of as the number of times the corresponding x coordinate occurs.

In other words, calculate $$\mbox{wmean}(x,y) = \frac{\sum_i x_i y_i}{\sum_i y_i}$$

Your function should check that the two arrays have the same size and shape, and return 0 if they don't. It should also return 0 if the sum of the y values is 0.


<!-- grade_id: cell-32b1fb0b42cd31e2, points: 0, grade: false, solution: true -->
```python
# wmean()
#YOUR CODE HERE
def wmean(arr1,arr2):
    if(not isinstance(arr1, np.ndarray) or not isinstance(arr2, np.ndarray) or arr1.ndim != 1 or arr2.ndim != 1):
                      return 0
    elif(arr1.size != arr2.size or arr1.shape != arr2.shape or arr2.sum() == 0):
                      return 0
    else: 
        return ((arr1*arr2).sum())/arr2.sum()
```

<!-- grade_id: cell-33bf073669925dc9, points: 1, grade: true, solution: false -->
```python
# test wmean() is defined
```

<!-- grade_id: cell-9a0f3a174b592755, points: 1, grade: true, solution: false -->
```python
# test wmean() calculations
```

<!-- grade_id: cell-8807ae578b4ca5e9, points: 1, grade: true, solution: false -->
```python
# Test wmean doesn't crash with invalid input
```

## Centering
Write a function, `center(text, length)` that will return a new string that contains the old string centered in a new string that has a length given by the length parameter. Your function should:
1. remove any whitepace (spaces, tabs etc) at the beginning and end of `text`
2. pad the new string with spaces before and after so that the text is centered in the new string.
3. If there are an odd number of spaces to add so that the text can not be perfectly centered, there should be one less space at the beginning. See the example test cases below.
4. If the string is longer than the specified length, your function should chop off charaters from the beginning and end of the string so that the returned string has the length specified. If an even number of characters need to be removed, then half should be removed from the beginning and half from the end. If an odd number of characters need to be removed, the extra character should be removed from the end. See the example test cases below.

Hint: look at the documentation for the built-in function `len()`.

<!-- grade_id: cell-e2a7f3bae1d7882e, points: 0, grade: false, solution: true -->
```python
#YOUR CODE HERE
def center(text, length):
    text = text.strip()
    spaces = length - len(text)
    
    if(len(text) > length):
        diff = len(text) - length
        
        if(diff%2 != 0):
            a = len(text)-1
            return text[(int)((diff-1)/2):(int)(a-(diff-1)/2)]                                
        else:
            text = (text[:-diff//2])
            text = (text[diff//2:])            
            return text
    
    elif(spaces%2 != 0):
        text = text.rjust(len(text) + (spaces-1)//2)
        text = text.ljust(len(text) + (spaces+1)//2)
    else:    
        text = text.rjust(len(text)+(spaces//2)) #left end
        text = text.ljust(len(text)+spaces//2) #right end
    
    return text
```

<!-- grade_id: cell-ff5f93beae9e9b57, points: 1, grade: true, solution: false -->
```python
# Test center() is defined
```

<!-- grade_id: cell-81faa53175185607, points: 0.5, grade: true, solution: false -->
```python
# Test cases with equal spaces on both sides:
assert(center("Hello!", 10) == "  Hello!  ")
assert(center("Hello!", 12) == "   Hello!   ")
```

<!-- grade_id: cell-e821d7da64da993c, points: 0.5, grade: true, solution: false -->
```python
# Test cases with unequal spaces
assert(center("Hello!", 11) == "  Hello!   ")
```

<!-- grade_id: cell-ceec8707d2323468, points: 0.5, grade: true, solution: false -->
```python
# Test that leading whitespace is removed
assert(center("   HI", 4) == " HI ")
assert(center("HI   ", 4) == " HI ")
```

<!-- grade_id: cell-e1b55b5b209332bd, points: 0.5, grade: true, solution: false -->
```python
# Test cases where length is shorter than string.
assert(center("432101234", 5) == "21012")
assert(center("432101234", 4) == "2101")
assert(center("     432101234 ", 5) == "21012")
```

Don't forget to:
1. Add the names, and describe the contributions of anybody besides yourself who contributed to your work
2. Ensure that your notebook runs without errors by selecting Validate and/or Kernel->Restart & Run All
3. Submit!