# Exercise 11

Oct 9, 2019. Due Oct 16, 9am.

Enter the names, and describe the contributions of anybody besides yourself who contributed/collaborated.


## Containers II

Write a function called `ends(t)` that takes a list as an argument and returns a tuple consisting of the first and last elements of the list. You may assume that t is a list with at least one element.

<!-- grade_id: cell-95817c9bc2fbadf4, points: 0, grade: false, solution: true -->
```python
def ends(t):
    if(len(t) == 1):
        c = (t[0],t[0])
        return c
    else:
        a = t[0]
        b = t[len(t)-1]
        c =(a,b)
    return c

```

<!-- grade_id: cell-Q1a, points: 1, grade: true, solution: false -->
```python
start, end = ends([1, 2, 3])
assert type(ends([1,2,3])) == tuple
```

<!-- grade_id: cell-4cb593dfdb388400, points: 1, grade: true, solution: false -->
```python
assert ends([1,3,4]) == (1,4)
assert ends([5.6, 'string', True, False, 7.3, [3, 4]]) == (5.6,[3, 4])
assert ends(['a']) == ('a', 'a')
```

## Broadcasting

Explain in a few words what each of the five commented expressions (which set variables e, f, g, h and i) in the cell below does, or, if it causes an error, explain why. An answer for e is given as an example.


```python
import numpy as np
a = np.arange(5)
b = np.random.random(4)
c = np.arange(20).reshape((4, 5))
d = b.reshape(4,1)

#e = a + b
#f = a + c
#g = b * c
#h = b * c.T
#i = a + d

```

<!-- grade_id: cell-fe9df6addfdb27fc, points: 2, grade: true, solution: true -->
e: this causes an error because a row vector of length 5 can't be added to a row vector length 4

f: this adds a to every array in inside b, so the 4 subarrays/row vectors in this case

g: this causes an error because b is a row vector of length 4 and you can't multiply it by c's 5 length row vectors

h: this transposes the matrix c and then multiples it by b

i: d is a column vector and we add every value in d to every value in a. This consequently creates a matrix with 5 columns and 4 rows. 

## Matrix multiplication with arrays

In this question you will verify that the numpy.ndarray matrix multiplication operator, @, does a proper matrix multiplcation. To do this, first write a function, `mat_mult(a1, a2)` that creates a new 2d array of zeros, and using nested for loops, fill in the elements of the matrix multiplication, and return this new array.

Then, write a second function, `matrix_diff(b, c)` that takes two arrays as arguments then generates and returns
the following ratio: $$ \sqrt{\frac{\sum (b_{i,j} - c_{i,j})^2 }{\sum (b_{i,j}+c_{i,j})^2} }$$ This function gives a measure of how different two arrays are, giving 0 if the arrays are identical.

Your functions can assume that both matrix arguments are two-dimensional, but they should check that the shapes are compatible and return None if they are not.

<!-- grade_id: cell-d79c9e7d26b8e43c, points: 0, grade: false, solution: true -->
```python
import numpy as np
def mat_mult(a1, a2):    
    cols1 = a1.shape[1]
    cols2 = a2.shape[1]
    rows1 = a1.shape[0]
    rows2 = a2.shape[0]
    if(cols1 != rows2):
        return None
   
    m = np.zeros((rows1,cols2))   
    for i in range(m.shape[0]): #row
        for j in range(m.shape[1]): #col
            for k in range(rows2):
                m[i][j] += (a1[i][k])*(a2[k][j])                
    return m

def matrix_diff(b, c):
    num = 0
    den = 0
    if(b.shape != c.shape):
        return None
    for i in range(b.shape[0]):
        for j in range(c.shape[1]):
            num += (b[i][j] -c[i][j])**2
            den += (b[i][j] +c[i][j])**2
    return np.sqrt(num/den)


```

<!-- grade_id: cell-de005f46ca5d284b, points: 1, grade: true, solution: false -->
```python
# check that mat_mult() gives the right answer:
from numpy.testing import assert_almost_equal, assert_equal
a = np.array([[1,2,3],[4,5,6]])
b = np.array([[1,5,9,1],[6,5,4,3],[1,2,3,4]])
assert_almost_equal( a@b, mat_mult(a, b), 6)
```

<!-- grade_id: cell-0fb906df9bbeddac, points: 1, grade: true, solution: false -->
```python
# check that matrix_diff gives 0 for arrays that are the same
cc = a@b
assert_almost_equal(matrix_diff(cc, cc), 0., 7)
```

<!-- grade_id: cell-da02b002cef30242, points: 1, grade: true, solution: false -->
```python
print(mat_mult(b,a))
if mat_mult(b, a) != None:
    raise AssertionError("call to mat_mult with incompatible arrays should return None")
if matrix_diff(a,b) != None:
    raise AssertionError("call to matrix_diff with different shape arrays should return None")
```

Output:
```
None
```

<!-- grade_id: cell-f5d227e37f9a0f36, points: 1, grade: true, solution: false -->
```python
# ensure matrix_diff gives the correct difference for matrices that aren't the same.
# and that mat_mult and @ give the same answer:
```

## Newton's Method III
Now we have all the tools we need to produce pictures similar to the one shown below
![](roots.png)
We won't get to making those plots today, but will work on optimizing our calculations in `which_root_z4()` to use array operations. Rewrite that function so that it takes an array of complex numbers as an argument. The argument will eventually be a 2d array whose values corresponds to the complex number of associated with a grid point. For example the first row of elements might have values: [0, 0.5, 1.0, 1.5] while the second row would then be: [0+0.5j, 0.5+0.5j, 1.0+0.5j, 1.5+0.5j] and the third row would be [0+1j, 0.5+1j. 1.0+1j, 1.5+1j]. Your function though should work with either 2d or 1d arrays of complex numbers as input. See the test cases below for some examples.

In your new `which_root_z4(z, n)` you should start by using a for loop to do n iterations of Newton's method for all points simultaneously. There should not be any loops over grid points, just array operations.

When the iterations are complete, create a new array, having one dimension more than the input array. The new array should have a shape of: (4, shape of input array). So if the input array is 2d with 10 rows and 20 columns, the new array should have a shape of (4, 10, 20). Use four lines of code to fill in the four planes of this array with measurements of the distances from each grid point to each of the four roots.

Then construct and return an array of the same size and shape as the input array, but of integers with values 0, 1, 2, or 3 indicating the root that the corresponding grid point has moved closest to (where the indices 0, 1, 2, and 3 correspond to the roots at 1, 1j, -1, -1j respectively).

All of this work should be done with no looping over grid points, use array operations. The only loop should be the for loop over the n iterations.

When calling `which_root_z4(z, n)` to test it, avoid giving a starting point of 0 anywhere, since the division by zero will cause trouble.


<!-- grade_id: cell-acf64714387c2139, points: 0, grade: false, solution: true -->
```python
def which_root_z4(z, n):
    roots = np.array([1,1j,-1,-1j])
    
    for i in range(n): #check if it should be n+1
        z_1 = z - f(z)/df(z)
        z = z_1 
    
    c = np.array([z,z,z,z]) #new array
    c[0] = abs(roots[0] - z)
    c[1] = abs(roots[1] - z)
    c[2] = abs(roots[2] - z)
    c[3] = abs(roots[3] - z)
    
    final = np.ndarray(shape = z.shape)
    final = np.argmin(c,axis = 0)
    
    return final

def f(z):
    return (z**4-1)
def df(z):
    return (4*z**3)
```

<!-- grade_id: cell-dc58d2a4ba6d371e, points: 1, grade: true, solution: false -->
```python
# test that the function executes and returns a 1d array of the right size with 1d array input
# here we just test a few different starting points
import numpy as np
from numpy.testing import assert_equal, assert_almost_equal
z1 = np.array([1.1j, 1.1, -2.5, -2j])
ind1 = which_root_z4(z1, 10)
assert ind1.ndim == 1
assert ind1.size == 4
#z = np.linspace(-)


```

<!-- grade_id: cell-ca7c24d26103dbfb, points: 1, grade: true, solution: false -->
```python
# test that the function returns a 2d array with 2d input:
# make an array of complex points that would be suitable for graphing:
x = np.linspace(-3, 3, 90) # misses 0, since it uses the end points. 91 would hit 0
x2d, y2d = np.meshgrid(x,x)
z = x2d + 1j*y2d
ind = which_root_z4(z, 10)
assert ind.ndim == 2
assert ind.size == 90*90
```

<!-- grade_id: cell-4c68bbfd63fdfc91, points: 1, grade: true, solution: false -->
```python
# check that in the 1d case we converge to the right values
assert_equal(ind1, [1, 0, 2, 3])

```

<!-- grade_id: cell-a994040db65e3978, points: 1, grade: true, solution: false -->
```python
# check a more complete array in 2d
```