# Exercise 10
Oct 7, 2019. Due Oct 9, 9am.

Enter the names, and describe the contributions of anybody **besides yourself** who contributed/collaborated. 

## Containers 
Below, you are asked to write several functions that manipulate lists and dictionaries.
Tests for most of these functions are visible. You should get in the habit of testing your own functions this way to ensure they do what is expected.


Write a function called `list2ends(t)` that takes a list as an argument and returns a list containing only the first and last elements of the argument. You can assume that the argument `t` is a list containing at least one element.



<!-- grade_id: cell-e81e4b8c84eaaa62, points: 0, grade: false, solution: true -->
```python
def list2ends(t):
    new_list = [t[0],t[len(t)-1]]
    if(len(t)==1):
        return t
    return new_list

```

<!-- grade_id: Q1a1, points: 1, grade: true, solution: false -->
```python
# check that function works:
assert list2ends([1, 3, 5]) == [1, 5]
assert list2ends(['b', 5, 'a string', 2]) == ['b', 2]
    
```

<!-- grade_id: Q1a2, points: 1, grade: true, solution: false -->
```python
# test edge cases
```

Write a function called `noduplicates(t)` that takes a list and returns a new list containing the elements of the first list but with any duplicates removed. The function should retain the first instance of any element but remove any later exact matches.

<!-- grade_id: cell-486df2d756bcb7f0, points: 0, grade: false, solution: true -->
```python
def noduplicates(t):
    new_list = []
    for n in range(len(t)):   
        if(new_list.count(t[n]) == 0):
            new_list.append(t[n])
    return new_list

```

<!-- grade_id: Q1b1, points: 1, grade: true, solution: false -->
```python
assert noduplicates([1, 3, 7, 'stuff', 45.2]) == [1, 3, 7, 'stuff', 45.2]
```

<!-- grade_id: Q1b2, points: 1, grade: true, solution: false -->
```python
assert noduplicates(['a', 'b', 4, 2, 9, 5.0, 'b']) == ['a', 'b', 4, 2, 9, 5.0]
assert noduplicates([1, 3, 3, 1, 12.9, 47.3, 'a', 3, 1, 'others', [4,2],'b',3]) == [1, 3, 12.9, 47.3, 'a', 'others', [4,2],'b']
```

Write a function called `lessthanten(t)` that takes a list of numbers and returns a new list containing only those elements with values less than 10, with no duplicates.

<!-- grade_id: cell-fc80db073cf0158e, points: 0, grade: false, solution: true -->
```python
def lessthanten(t):
    lt =  noduplicates(t)
    new_l = []
    for i in range(len(lt)):
        if(lt[i]<10):
            new_l.append(lt[i])
    return new_l
```

<!-- grade_id: Q1c1, points: 1, grade: true, solution: false -->
```python
assert lessthanten([1, 2, 7, 38, 5.2, -27.4,  24]) == [1, 2, 7, 5.2, -27.4]
```

<!-- grade_id: Q1c2, points: 1, grade: true, solution: false -->
```python
assert lessthanten([1, 2, 7, 38, 5.2, -27.4, 7, 1, 5.2, 24]) == [1, 2, 7, 5.2, -27.4]
```

Write a function called `truncate(t)` that takes a list `t`, removes the first and last elements, modifying it in place. The function should return None. Your function should not crash if it is passed a list with only 1, or with no elements.

<!-- grade_id: cell-65168bfb5eae051e, points: 0, grade: false, solution: true -->
```python
def truncate(t):
    if(len(t) == 0):
        return None
    elif(len(t) == 1):
        t.pop()
    else:
        t.pop()
        t.remove(t[0])
        
    return None
        
```

<!-- grade_id: Q1d1, points: 1, grade: true, solution: false -->
```python
a = [9, 8, 7, 6]
assert truncate(a) == None
assert a == [8, 7]
```

<!-- grade_id: Q1d2, points: 1, grade: true, solution: false -->
```python
a = [4]
assert truncate(a) == None
assert a == []
a = [ 5, 2]
assert truncate(a) == None
assert a == []
a = ['234', 234, True, 24.2]
assert truncate(a) == None
assert a == [234, True]
```

Write a function `concatenate(a, b, c)` which returns the concatenated dictionaries a, b, c, but leaves a, b, and c untouched.

<!-- grade_id: cell-6142d95bd336bb13, points: 0, grade: false, solution: true -->
```python
import numpy as np
def concatenate(a,b,c):
    d = {}
    d.update(a)
    d.update(b)
    d.update(c)
    return d
```

<!-- grade_id: Q1e1, points: 1, grade: true, solution: false -->
```python
a = {'type': 'potato', 'colour': 'brown', 'shape': 'ellipsoid'}
b = {'a': 'alpha', 'b': 'beta'}
c = {'name1': 'charlie', 'height': 123}

d = concatenate(a, b, c)
assert d == {'type': 'potato', 'colour': 'brown', 'shape': 'ellipsoid', 'a': 'alpha', 'b': 'beta', 'name1': 'charlie', 'height': 123}

```

<!-- grade_id: Q1e2, points: 1, grade: true, solution: false -->
```python
assert a == {'type': 'potato', 'colour': 'brown', 'shape': 'ellipsoid'}
assert b == {'a': 'alpha', 'b': 'beta'}
assert c == {'name1': 'charlie', 'height': 123}
```

Write a function `exists(d, s)` which checks if key s already exists in dictionary d. It should return a boolean value of True if the key s is in d.

<!-- grade_id: cell-a1af248c383a6e32, points: 0, grade: false, solution: true -->
```python
def exists(d,s):
    keys = list(d.keys())
    for n in range(len(keys)):
        if(s == keys[n]):
            return True
    return False

```

<!-- grade_id: Q1f1, points: 1, grade: true, solution: false -->
```python
# check if exists() returns the correct value if s is not in d
a = {'one': 'potato', 'red': 'sock', 'alpha': 'a', 'beta': 'b', 'name1': 'charlie', 'name2': 'chaplin'}
assert exists(a, 'key') == False
```

<!-- grade_id: Q1f2, points: 1, grade: true, solution: false -->
```python
# check if exists returns the correct value if s is not in d
a = {'one': 'potato', 'red': 'sock', 'alpha': 'a', 'beta': 'b', 'name1': 'charlie', 'name2': 'chaplin'}
assert exists(a, 'one') == True
```

## Numerical Integration I

The simplest numerical integration technique uses the midpoint rule. The definite integral of a
function $f(x)$ over the interval $x = [x_0, x_n ]$ is approximated as a sum of $n$ rectangle areas of equal
width $\Delta x = (x_n −x_0)/n$ and height determined by the value of the function $f (x^*_i)$ at the midpoint
position $x^*_i$ . This is illustrated in the figure. ![](midpoint.png)

Mathematically the midpoint rule estimate for the integral is given by:
$$\int_{x_0}^{x_n} f(x) \, dx \approx \Delta x  \, \left(f(x_1^*) + f(x_2^*) + ... + f(x_{n}^*) \right)$$

Write a function called `midpoint(n, x0, xn)` that calculates the integral of $f(x) = \sin(x) \exp(-x/2)$ over the interval from `x0` to `xn` using `n` intervals. 

<!-- grade_id: cell-b7cee818a10f64ee, points: 0, grade: false, solution: true -->
```python
import math

def midpoint(n,x0,xn):
    width = (xn-x0)/n
    ans = 0;
    x = x0+(width/2)
    
    for i in range(n):
        ans += width*(math.sin(x)*math.exp((-x)/2))
        x += width
    return ans


```

<!-- grade_id: Q2a, points: 1, grade: true, solution: false -->
```python
# check that the function exists and returns a floating point number
a = midpoint(20, 0., 1.)

import numbers
from numpy.testing import assert_equal, assert_almost_equal
assert isinstance(a, numbers.Real) == True
assert isinstance(a, numbers.Integral) == False
```

```python
# check that the value is correct:
assert_almost_equal( midpoint(200, 0., 1.) , 0.333681854288, 5)

```

<!-- grade_id: Q2b, points: 1, grade: true, solution: false -->
```python
# check that the correct number of steps are used

```

Use your `midpoint` function to calculate the integral of `f(x)` over the interval $ x = [ -0.5, 2\pi]$ for values of `n = [10, 20, 100, 200, 400, 800, 1600]` Store the values in a list called `integrals`.

<!-- grade_id: cell-7d0a58ed6f9a6d94, points: 0, grade: false, solution: true -->
```python
n = [10, 20, 100, 200, 400, 800, 1600]
integrals = []

for i in range(len(n)):
    val = midpoint(n[i],-0.5,math.pi)
    integrals.append(val)

```

<!-- grade_id: Q2c, points: 1, grade: true, solution: false -->
```python
assert_almost_equal(np.array(integrals[2:]), np.array([0.8216272417353141, 0.8215591715038066, 0.8215421543892597, 0.8215379001383317, 0.8215368365773316]),5)
```

The indefinite integral of $f(x)$ is $F(x) = -\frac{2}{5}\exp(-x/2) (\sin(x) +2\cos(x)) + \mbox{const.}$

Make a plot of your midpoint integrals vs n, using red circles for the points. On the same plot, indicate the exact value of the integral with a horizontal dashed line. Don't forget axes labels and a title. Use a legend box to indicate the exact value and the midpoint calculations. Add a text label to the plot, just above the dashed line, near the right edge that gives the exact answer, to 6 decimal places.


<!-- grade_id: cell-855a687fd623b451, points: 2, grade: true, solution: true -->
```python
import matplotlib.pyplot as plt

def inter(x):
    return -0.4*math.exp(-x/2)*(math.sin(x)+2*math.cos(x))

rval = inter(math.pi)-inter(-0.5)
plt.plot(n,integrals, 'ro', label = 'Midpoint Values')
plt.axhline(y=rval,ls = '--', label  = 'Exact Value')
plt.xlabel('n value')
plt.ylabel('Integral values')
plt.title('Midpoint estimates vs Indefinite integral')
plt.text(1600-200,rval+0.0004,'{0:.6f}'.format(rval))
plt.legend()
plt.show()
```

Output:
```
<Figure size 640x480 with 1 Axes>
```

Finally, write a function `find1(n, integrals, exact)` that will take three arguments: a list of n values, a list of integrals calculated as above, and an exact value. The function should find the lowest n value in the list for which the corresponding integral is within 0.1%. It should return this lowest value of n. If none of the integral values are within 0.1% of the exact value, the function should return 0. The function may assume that the arguments are legal, and that the lists of n and integral values have then same number of elements, but it should not assume that the n values have any particular order.

<!-- grade_id: cell-d00172b787982e1c, points: 0, grade: false, solution: true -->
```python
def find1(n, integrals, exact):
    best_n = -1
    
    for i in range(len(integrals)):       
        if(abs(integrals[i]-exact)/exact < 0.001):
          
            if(best_n == -1):
                best_n = n[i]
            else:                
                if(n[i]<best_n):
                    best_n = n[i]
                                       
    return best_n

```

<!-- grade_id: cell-a578a511e6d3b604, points: 1, grade: true, solution: false -->
```python
n = [10, 20, 100, 200, 400, 800, 1600]
integrals = [0.8306373587315234, 0.8238069863597768, 0.8216272417353141, 0.8215591715038066, 0.8215421543892597, 0.8215379001383317, 0.8215368365773316]
ans = 0.8215364820571522
assert find1(n, integrals, ans ) == 100

```

<!-- grade_id: cell-774440ce5efed952, points: 1, grade: true, solution: false -->
```python
# check trickier cases
```