## In-class assignment 2
1. You are free to access course materials and the internet, but no communication with other people (either in person or electronically) NO PHONES.
2. Please Save and Submit promptly when the time is up. There will be a severe late penalty: grades will be scaled by a decaying exponential with a time constant of 3 minutes.


Follow the instructions for Question 1: (3 points)

<!-- grade_id: cell-d339f205d34ed7ab, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    ans = 1
    if(n == 0 or n==1):
        return ans
    else:
        for i in range(2,n+1):
            ans = i*ans            
        return ans    
    
```

<!-- grade_id: cell-62d44efddfc1e0fe, points: 1, grade: true, solution: false -->
```python
# check that function executes and returns an integer
```

<!-- grade_id: cell-9d6b7380f0dab4e2, points: 1, grade: true, solution: false -->
```python
# check that return values are correct
```

<!-- grade_id: cell-55de154f2b95b55a, points: 1, grade: true, solution: false -->
```python
# check that return values are correct
```

Follow the instructions for Question 2 here: (3 points)

<!-- grade_id: cell-34082f06e31cf236, points: 0, grade: false, solution: true -->
```python
import numpy as np

def mult(a):
    for i in range(0,len(a)):
        a[i] = a[i]*i
    return a

```

Output:
```
[ 0  4 10 18 28]
```

<!-- grade_id: cell-022946b94cf689bf, points: 1, grade: true, solution: false -->
```python
# check that the function executes
```

<!-- grade_id: cell-ee104da8a45d972c, points: 1, grade: true, solution: false -->
```python
# check that it returns an array of the correct dimensions
```

<!-- grade_id: cell-f69ed7bb0c885153, points: 1, grade: true, solution: false -->
```python
# check return values
```

Follow the instructions for Question 3 here. (6 points)

<!-- grade_id: cell-05e0d194f3d50179, points: 0, grade: false, solution: true -->
```python
import numpy as n
def rewrite():
    f = open('names.txt','r')
    text = f.readlines()
    new_f = open('newnames.txt','w')
    arr = []
    for i in range(len(text)):        
        if(not text[i].startswith('#')):  
            arr = text[i].split(',')
            arr[1] = arr[1].strip()
            new_f.write("{}'s number is {}.\n".format(arr[1], arr[0])) 
            if(i == len(text)-1):
                print(arr[1])
    new_f.close()
    return text
        

```

<!-- grade_id: cell-e98fdee7476449c2, points: 1, grade: true, solution: false -->
```python
# check that rewrite returns a list of lines
a = rewrite()
assert type(a) == list
assert len(a) == 7

```

Output:
```
Mark Watney
```

<!-- grade_id: cell-7e8ca9ed495cbcf7, points: 1, grade: true, solution: false -->
```python
# check return value

```

<!-- grade_id: cell-f5f9a76de4139a9a, points: 1, grade: true, solution: false -->
```python
# check the file is created
```

<!-- grade_id: cell-ebe85f713f773106, points: 1, grade: true, solution: false -->
```python
# check that the content is reasonable
```

<!-- grade_id: cell-3520c255e713ee9f, points: 1, grade: true, solution: false -->
```python
# check that the content is correct
```

<!-- grade_id: cell-8038af802dd95aa4, points: 1, grade: true, solution: false -->
```python
# check that the content is correct
```